# Fonctions logiques


# Utilisation des fonctions logiques en python

Python posséde les opérateurs suivants:

|opérateur|symbolique|syntaxe Python|
|:-|:-:|-:|
|NON|$\bar{a}$|`not(a)`|
|ET|$a\times b$|`a and b`|
|OU|$a + b$|`a or b`|
|XOR|$a\oplus b$|non implémentée|

Quelques exemples :

```python
a = True
b = False
print("1 : ", not(a))
print("2 : ", a and b)
print("3 : ", a or b)
```

Output:
```
1 :  False
2 :  False
3 :  True
```

Cependant, pour utiliser le OU EXCLUSIF (xor) , il faut utiliser soit une bibliothèque complémentaire, soit définir soit même sa fonction. 
C'est donc ce que vous allez devoir faire : definir une fonction `xor(a,b)` qui prend deux boolèns en arguments en renvoie un boolèen correspondant à la table de vérité du XOR :

|  a  |   b   |   S   |
|:-:|:-:|:-:|
|0|0|0|
|0|1|1|
|1|0|1|
|1|1|0|

On peut pour cela utiliser la relation $a\oplus b = (a\times \bar{b}) + (\bar{a} \times b)$

<!-- grade_id: cell-ab182df9e6ff4415, points: 0, grade: false, solution: true -->
```python
def XOR(a,b) :
    ### BEGIN SOLUTION
    x = (a and not(b)) or (not(a) and b)
    return(x)
    ### END SOLUTION
```

Testons maintenant notre fonction pour qu'elle réponde bien à la table de vérité du OU EXCLUSIF :

<!-- grade_id: cell-063b493fafc28407, points: 1, grade: true, solution: false -->
```python
assert XOR(False,False)==False
assert XOR(False,True)==True
assert XOR(True,False)==True
assert XOR(True,True)==False
```

## 2. Egalité de fonctions logiques
On souhaite tester une relation logique appelée **théorème de Morgan** qui dit :

$$\bar{A+B}=\bar{A} \times \bar{B}$$

Pour cela vous devez créer deux fonctions :
- `logique1(A,B)`, qui correspond au premier terme de l'égalité du théorème de Morgan 
- `logique2(A,B)`, qui correspond au second terme de l'égalité du théorème de Morgan.

Les tests unitaires testerons ensuite l'égalité des deux fonctions.

<!-- grade_id: cell-4a8e1435b1043fa4, points: 0, grade: false, solution: true -->
```python
def logique1(A,B) :
    ### BEGIN SOLUTION
    x = not(A or B)
    return(x)
    ### END SOLUTION
    
def logique2(A,B) :
    ### BEGIN SOLUTION
    x = not(A) and not(B)
    return (x)
    ### END SOLUTION
```

effectuons maintenant les tests qui permettrons de vérifier l'égalité des deux fonctions :

<!-- grade_id: cell-7c0e08253af31468, points: 2, grade: true, solution: false -->
```python
assert logique1(False,False)==logique2(False,False)
assert logique1(False,True)==logique2(False,True)
assert logique1(True,False)==logique2(True,False)
assert logique1(True,True)==logique2(True,True)
```

Remarque : ceci n'est qu'un exercice, vous n'avez pas à connaitre le théorème de Morgan.